![Banner](img/banner.png)

# **Actividad N°3:** Soluciones básicas

***Matemáticas para Machine Learning***.

**Semana 1 - Lección 3:** Sistemas de ecuaciones lineales.

**Profesor:** *Fernando Lozano* - **Autor Notebook:** *César Garrido Urbano*

# Introducción

## Descripción




El presente *jupyter notebook* contiene todo el material para el desarrollo de la Actividad 3 de la Semana 1 del curso ***Matemáticas para Machine Learning***. En este se realizarán ejercicios numéricos relacionados con la búsqueda de soluciones básicas dentro de un sistema de ecuaciones lineales y se pondrán en práctica los conceptos vistos en la Lección 3.

**Objetivos de Aprendizaje:**

*   Determinar si un sistema de ecuaciones tiene solución única, soluciones infinitas o no tiene soluciones.
*   Entender e implementar el proceso para encontrar soluciones básicas dentro de un sistema lineal de la forma $Ax = b$.
*   Automatizar el proceso para encontrar soluciones básicas a partir de la implementación de una función en Python.

## Teoría



Considere un sistema de $m$ ecuaciones con $n$ incógnitas.

$$ m \left\{ \begin{split}
  a_{11} x_1 + a_{12} x_2 + \ldots + a_{1n} x_n & = b_1 \\ 
  a_{21} x_1 + a_{22} x_2 + \ldots + a_{2n} x_n & = b_2 \\ 
  \vdots \qquad \qquad \quad & \\
  a_{m1} x_1 + a_{m2} x_2 + \ldots + a_{mn} x_n & = b_m \\ 
 \end{split} \right. $$

 Y suponga que tanto sus columnas como sus filas son linealmente independientes entre sí. Entonces tenemos 3 casos:

*   $m \leq n$ : Tenemos menos ecuaciones que incógnitas, y debido a la independencia lineal el sistema no tiene soluciones 
*   $ m = n$ : Tenemos la misma cantidad de ecuaciones que de incógnitas, de nuevo por independencia lineal tenemos una única solución
*   $ m \geq n$ : Tenemos mas ecuaciones que incógnitas, y debido a independencia lineal tenemos infinitas soluciones. Estas infinitas soluciones forman un espacio vectorial de dimensión $m-n$

El anterior sistema también se puede escribir de la siguiente manera:

\begin{equation}
      A \>\mathbf{x} =\mathbf{b} \qquad A \in \mathbb{R}^{m \times n} ,\mathbf{x} \in \mathbb{R}^n,\mathbf{b} \in \mathbb{R}^m 
\end{equation}


\begin{equation} A = \begin{bmatrix} 
    a_{11} & a_{12} & \ldots & a_{1n} \\
    a_{21} & a_{22} & \ldots & a_{2n} \\
    \vdots & \vdots & \ddots & \vdots \\
    a_{m1} & a_{m2} & \ldots & a_{mn} \\
\end{bmatrix} \end{equation}

## Metodología

De forma específica, ahora considere de un sistema de $m$ ecuaciones y $n$ incógnitas en donde $m \leq n$:

\begin{equation}
      A \>\mathbf{x} =\mathbf{b} \qquad A \in \mathbb{R}^{m \times n} ,\mathbf{x} \in \mathbb{R}^n,\mathbf{b} \in \mathbb{R}^m 
\end{equation}

En este caso, una solución básica del sistema de ecuaciones es aquella en la cual a lo sumo hay $m$ variables diferentes de cero.

Encontrar estas soluciones serán el objetivo de esta actividad!

El siguiente *jupyter notebook* consiste de 3 partes: 

0.   Un ejemplo detallado del procedimiento para encontrar las soluciones básicas del sistema anteriormente presentado.
1.   Unos ejercicios para que ponga a prueba sus habilidades resolviendo 3 distintos sistemas de ecuaciones propuestos.
2.   Un último punto donde tendrá que completar el bosquejo de una función que realiza todo el proceso de forma automatica.


Para esto usted debe seguir siempre el mismo formato. Al encontrar/guardar las soluciones básicas del sistema utilice un diccionario:
    
$$
    \texttt{sol = \{ $'$Soluciones$'$ : [ ], $'$Bases$'$ = [ ]      \} }
$$

En donde cada elemento de $\texttt{dir[ $'$Soluciones$'$ ]}$ es un vector en $\mathbb{R}^n$ con por lo menos $n-m$ variables iguales a 0 (soluciones básicas) y cada elemento de $\texttt{dir[ $'$Bases$'$ ]}$ contiene una matriz en $\mathbb{R}^{m \times n}$ con la base correspondiente a dicha solución básica. Cada matriz tiene por lo menos $n-m$ columnas iguales a 0 (aquellas que no hacen parte de la base).

Utilice la función $\texttt{pretty_print()}$ para imprimir el resultado en consola (véase el ejemplo). De esta manera podrá visualizar de  mejor forma los datos.

Para desarrollar el taller usted deberá editar las celdas de código dispuestas para esto. Estas estarán marcadas con el siguiente comentario:

```python
# =====================================================
# COMPLETAR ===========================================
# 

# =====================================================
```

Edite o complete el códgio dentro de estas lineas de comentarios. Dentro de estos comentarios encontrará indicaciónes de lo que debe hacer, así como algunas de las variables que debe utilizar o calcular (puede que estas tengan ya una estructura para llenar o no, revise y complete la asignación).

# Ejemplos de Soluciones básicas

## Inicialización

Se importan librerías básicas para el desarrollo del taller y se define la función $\texttt{pretty_print(df)}$ para la visualización de los resultados de forma más adecuada.

```python
# Librerias principales
import numpy as np
import pandas as pd
import sympy
import copy
import sys

# Visualización
from IPython.display import display, HTML
import matplotlib.pyplot as plt 
import matplotlib.colors as mcolors

# Combinaciones
from itertools import combinations

# Maia Utils
sys.path.append('../')
from MaiaUtils_S1TP3 import MaiaUtils
```

```python
# Maia Utis
maia = MaiaUtils()

# Aumentar tamaño de columna DataFrame
pd.set_option('max_colwidth', 80)

# Crear función para ocultar seguimiento de debugging
if "toggle_traceback" not in globals():
    ipython = get_ipython()
    toggle_traceback = maia.manipulate_ipython(ipython)
```

```python
# Cambiar configuración de debugging
toggle_traceback()
```

Considere la siguiente matriz de ejemplo:

$$ A = \begin{bmatrix}
  1 & 2 & 3 \\
  2 & 3 & 6 
\end{bmatrix} 
\qquad \mathbf{b} = \begin{bmatrix}
  4 \\ 8
\end{bmatrix} \qquad \mathbf{x} = \begin{bmatrix}
x_1 \\ x_2 \\ x_3
\end{bmatrix}$$

## Solución 1: Columnas 1 y 2

Tome las columnas 1 y 2 como base y resuelva:

$$ 
  \begin{bmatrix} 1 & 2 & 0 \\ 2 & 3 & 0 \end{bmatrix} 
  \begin{bmatrix} x_1 \\ x_2 \\ 0 \end{bmatrix} 
  = \begin{bmatrix} 4 \\ 8 \end{bmatrix} 
  \quad ⟶ \quad 
  \begin{array}{rl} x_1 = & 4 \\ x_2 = & 0 \\ x_3 = & 0  \end{array}
$$

```python
# Se define la submatriz A1 con únicamente las columnas 1 y 2 como base
A1 = [[1, 2],
      [2, 3]]
# Vector b
b = [4, 8]

# Se resuelve utilizando la función np.linalg.solve de numpy
x1 = np.linalg.solve(A1, b)

# Se imprime el resultado
print(f'x1 = {x1}')
```

## Solución 2: Columnas 1 y 3

Tome las columnas 1 y 3 como base:

$$ \begin{bmatrix} 1 & 0 & 3 \\ 2 & 0 & 6 \end{bmatrix} 
   \begin{bmatrix} x_1 \\ 0 \\ x_3 \end{bmatrix} 
   = \begin{bmatrix} 4 \\ 8 \end{bmatrix} $$





```python
# Se define la submatriz A1 con únicamente las columnas 1 y 2 como base
A2 = [[1, 3],
      [2, 6]]

# Se intenta resolver, pero el código falla
try:
    x2 = np.linalg.solve(A2, b) 
# Se imprime el Error -> Matriz singular
except Exception as e:
    print(f'Error: {e}')
```

Observe que en este caso la primera y segunda fila de la matriz son linealmente dependientes. Esto causa el error de Matriz Singular al intentar resolver el sistema. Note que:

$$ 2\left( \begin{bmatrix} 1 & 0 & 3 \end{bmatrix} \begin{bmatrix} x_1 \\ 0 \\ x_3 \end{bmatrix} = 4 \right) 
\quad \sim \quad 
\left( \begin{bmatrix} 2 & 0 & 6 \end{bmatrix} \begin{bmatrix} x_1 \\ 0 \\ x_3 \end{bmatrix} = 8 \right) $$

Sin embargo, este sistema (de dos ecuaciones) se puede reducir en otro equivalente (de una ecuación) con 2 soluciones básicas adicionales:

$$ \begin{bmatrix} 1 & 0 & 3 \end{bmatrix} 
\begin{bmatrix} x_1 \\ 0 \\ x_3 \end{bmatrix} 
= \begin{bmatrix} 4 \end{bmatrix} 
\quad ⟶ \quad 
\begin{split}
& \begin{bmatrix} 1 & 0 & 0 \end{bmatrix} 
\begin{bmatrix} x_1 \\ 0 \\ 0 \end{bmatrix} 
= \begin{bmatrix} 4 \end{bmatrix} 
\quad ⟶ \quad 
\begin{array}{rl} x_1 = & 4 \\ x_2 = & 0 \\ x_3 = & 0  \end{array}
\\ \\
& \begin{bmatrix} 0 & 0 & 3 \end{bmatrix} 
\begin{bmatrix} 0 \\ 0 \\ x_3 \end{bmatrix} 
= \begin{bmatrix} 4 \end{bmatrix} 
\quad ⟶ \quad 
\begin{array}{rl} x_1 = & 0 \\ x_2 = & 0 \\ x_3 = & 4/3  \end{array}
\end{split} $$

```python
# Se guardan las soluciones
x21 = np.array([4.0, 0.0 ])
x22 = np.array([0.0, round(4/3,2)])

print(f'x21 = {x21}')
print(f'x22 = {x22}')
```

### Teoria adicional

Tome como ejemplo la siguiente matriz $A$ $(2\times n)$, con filas $a_1$ y $a_2$ y el vector $b$ con valores $b_1$ y $b_2$.

$$ A = \begin{bmatrix} - & \mathbf{a}_1 & - \\ - & \mathbf{a}_2 & -  \end{bmatrix} \qquad \mathbf{a}_1,\mathbf{a}_2 \in \mathbb{R}^3 \qquad \qquad \mathbf{b} = (b_1, b_2) \qquad b_1, b_2 \in \mathbb{R}$$

En caso de que $\mathbf{a}_1 = \alpha \> \mathbf{a}_2$, se cumple que:

$$ \begin{split} \mathbf{b}_1 \neq \alpha \> \mathbf{b}_2 \quad & \longrightarrow \quad \textrm{El sistema no tiene solución} \\
\mathbf{b}_1 = \alpha \> \mathbf{b}_2 \quad & \longrightarrow \quad \textrm{El sistema tiene 2 soluciones básicas} 
\end{split} $$

De manera general, cuando trabajamos con un sistema de ecuaciones identificado con una matriz $A \in \mathbb{R}^{m \times n}, n \geq m$ , que tiene $n-m$ columnas en $0$, es equivalente a trabajar con el sistema de ecuaciones identificado con la matriz $A' \in \mathbb{R}^{m \times m}$ que solo mantiene las columnas que són distintas de $0$, entonces el subsitema que se obtiene es:

$$ \mathbf{b} = (b_1, \ldots, b_m) \quad  A = \begin{bmatrix} - & \mathbf{a}_1 & - \\ \vdots & \vdots & \vdots  \\ - & \mathbf{a}_m & -  \end{bmatrix} \qquad \mathbf{a}_1, \ldots ,\mathbf{a}_m \in \mathbb{R}^m$$

si $a_k$ es combinación lineal de los otros $a_i$'s de tal manera que:

$$ a_k = \sum_{i \neq k} \alpha_i \mathbf{a}_i $$

entonces:

$$ \begin{split} \mathbf{b}_1 \neq \sum_{i \neq k} \alpha_i \> \mathbf{b}_i \quad & \longrightarrow \quad \textrm{El sistema no tiene solución} \\
\mathbf{b}_1 = \sum_{i \neq k} \alpha_i \> \mathbf{b}_i \quad & \longrightarrow \quad \textrm{El sistema tiene por lo menos } \left(\begin{array}{c} m \\ m-1 \end{array}\right) \textrm{ soluciones básicas} 
\end{split} $$

En este caso el número de soluciones básicas aumenta a partir del número de subsistemas que podemos generar una vez eliminamos la ecuación linealmente independiente de la matriz, obteniendo una matrix $(m-1) \times m $. Esto es $m$ combinado $m-1$

## Solución 3: Columnas 2 y 3

   $$ \begin{bmatrix} 0 & 2 & 3 \\ 0 & 3 & 6 \end{bmatrix} 
   \begin{bmatrix} 0 \\ x_2 \\ x_3 \end{bmatrix} 
   = \begin{bmatrix} 4 \\ 12 \end{bmatrix} 
   \quad ⟶ \quad 
   \begin{array}{rl} x_1 = & 0 \\ x_2 = & -4 \\ x_3 = & 4  \end{array}$$


```python
# Se define la submatriz A3 con únicamente las columnas 2 y 3 como base
A3 = [[2, 3],
      [3, 6]]

# Se resuelve
x3 = np.linalg.solve(A3, b)
x3 = np.round(x3,2)
print(f'x3 = {x3}')
```

## Visualización de las soluciones


Ahora se compilan todas las soluciones en el formato propuesto (diccionario) y se visualizan con la función $\texttt{pretty_print()}$.

Recuerde que en total se encontraron 4 soluciones. Esto por lo que cuando se tomaron las columnas 1 y 3 como base, se obtuvieron 2 soluciones.

```python
# Diccionario de la solución
sol_ejemplo = {'Soluciones':[], 'Bases':[]}

# Matriz A de ejemplo
A = np.array([[1, 2, 3],
              [2, 3, 6]])

# Columnas tomadas para cada solución (Recuerde que python es cero indexado)
col_idx = [[0,1], [0,2], [0,2], [1,2]] 
# Vector con soluciones de los sistemas reducidos
sol_reducidas = [x1, x21, x22, x3]

# Llenar diccionario de las soluciones guardadas y con las bases correspondientes
for col, x in zip(col_idx, sol_reducidas):
    # Base con columna i-ésima en 0
    A_base = A.copy()
    A_base[:,[i for i in range(A.shape[1]) if i not in col]] = 0
    # Solución con coordenada i'esima 0
    x_sol = np.zeros((3))
    x_sol[col] = x

    # Se guarda la solucion y la base en el diccionario
    sol_ejemplo['Soluciones'].append(x_sol)
    sol_ejemplo['Bases'].append(A_base)
```

```python
# Se imprime en consola con la función pretty_print para verla en forma de tabla
maia.pretty_print(sol_ejemplo)
```

# Problema 1: Soluciones básicas

Ahora usted debe realizar el mismo proceso descrito anteriormente para cada una de los sistemas propuestos:

## Matriz a)

 $$ A_a = 
 \begin{bmatrix} 1 & 2 & 3 & 4 \\ 2 & 0 & 4 & 0 \end{bmatrix} 
 \quad  \mathbf{b_a} = 
 \begin{bmatrix} 1 \\  2 \end{bmatrix} 
 \qquad \mathbf{x_a} = 
 \begin{bmatrix} x_1 \\  x_2 \\ x_3 \\ x_4 \end{bmatrix} $$

```python
# Diccionario de la solución a)
sol_a = {'Soluciones':[], 'Bases':[]}

# Matriz A y vector b del punto a)
Aa = np.array([[1, 2, 3, 4],
               [2, 0, 4, 0]])
ba = np.array([1, 2])

# ESTUDIANTE:
# Complete el diccionario 'sol_a' con las soluciones básicas y sus bases correspondientes.

# Col 0 y 1
Aa_01 = [[1, 2],
         [2, 0]]
x_01 = np.linalg.solve(Aa_01, ba)

# Col 0 y 2
Aa_02 = [[1, 3],
         [2, 4]]  
x_02 = np.linalg.solve(Aa_02, ba)

# Col 0 y 3
Aa_03 = [[1, 4],
         [2, 0]]  
x_03 = np.linalg.solve(Aa_03, ba)

# Col 1 y 3 (Submatriz singular)
Aa_13 = [[2, 4],
         [0, 0]]  
# Note que al intentar solucionar esta submatriz el solver falla,
# por esta razón se guarda el error (string) como constancia del resultado.
try:
    x_13 = np.linalg.solve(Aa_13, ba)
except Exception as e:
    x_13 = str(e)

# =====================================================
# COMPLETAR ===========================================
# -
# Completar el diccionario 'sol_a' con las soluciones básicas faltantes.
# -





# Columnas usadas como base en cada solución (Recuerde que python es cero indexado)
col_idx = [[0,1], [0,2], [0,3], [1,3]] 
# Vector con soluciones de los sistemas reducidos
sol_reducidas = [x_01, x_02, x_03, x_13]

# =====================================================

# Llenar diccionario de las soluciones guardadas y con las bases correspondientes
for col, x in zip(col_idx, sol_reducidas):
    # Base: Matriz A completa con columnas i'esimas (col_idx) en 0
    A_base = Aa.copy()
    A_base[:,[i for i in range(Aa.shape[1]) if i not in col]] = 0
    # Solución: Vector 4x1 con coordenadas i'esimas (col_idx) en 0
    if type(x) is np.ndarray:
        x_sol = np.zeros(Aa.shape[1])
        x_sol[col] = x
    # No se altera en caso de que sea el resultado guardado de la matriz singular
    else:
        x_sol = x

    # Se guarda la solucion y la base en el diccionario
    sol_a['Soluciones'].append(x_sol)
    sol_a['Bases'].append(A_base)

```

```python
maia.pretty_print(sol_a)
```

## Matriz b)

 $$  A_b = 
 \begin{bmatrix} 
 1 & 1 & 1 & 1 & 1 \\ 1 & 2 & 3 & 4 & 5 \\ 4 & 3 & 3 & 3 & 4  
 \end{bmatrix}
  \quad  \mathbf{b_b} = 
 \begin{bmatrix} 1 \\  2 \\ 3 \end{bmatrix} 
 \qquad \mathbf{x_b} = 
 \begin{bmatrix} x_1 \\  x_2 \\ x_3 \\ x_4 \\ x_5 \end{bmatrix}  $$

```python
# Diccionario de la solución b)
sol_b = {'Soluciones':[], 'Bases':[]}

# Matriz A y vector b del punto b)
Ab = np.array([[1, 1, 1, 1, 1],
              [1, 2, 3, 4, 5],
              [4, 3, 3, 3, 4]])
bb = np.array([1, 2, 3])

# =====================================================
# COMPLETAR ===========================================
# -
# Completar el diccionario 'sol_b' con las soluciones básicas.
# -





# =====================================================
```

```python
maia.pretty_print(sol_b)
```

## Matriz c)

 $$  A_c = 
 \begin{bmatrix} 
 1 & 1 & 1 & 1 & 1 \\ 1 & 2 & 3 & 4 & 5 \\ 2 & 3 & 4 & 5 & 6
 \end{bmatrix}
  \quad  \mathbf{b_c} = 
 \begin{bmatrix} 1 \\  2 \\ 3 \end{bmatrix} 
 \qquad \mathbf{x_c} = 
 \begin{bmatrix} x_1 \\  x_2 \\ x_3 \\ x_4 \\ x_5 \end{bmatrix}  $$

```python
# Diccionario de la solución c)
sol_c = {'Soluciones':[], 'Bases':[]}

# Matriz A y vector b del punto c)
Ac = np.array([[1, 1, 1, 1, 1],
               [1, 2, 3, 4, 5],
               [2, 3, 4, 5, 6]])
bc = np.array([1, 2, 3])

# =====================================================
# COMPLETAR ===========================================
# -
# Completar el diccionario 'sol_c' con las soluciones básicas.
# -





# =====================================================
```

```python
maia.pretty_print(sol_c)
```

## Calificación P1

<!-- grade_id: cell-a935c95b15fe4e3e, points: 10, grade: true, solution: false -->
```python
# Verificación de formato (10pts)
for i, sol in zip(["a", "b", "c"], [sol_a, sol_b, sol_c]):
    
    # Tipo de la solución
    assert type(sol) == dict, "Ha cambiado el tipo de la solución, este debe ser un diccionario!"
    for s in sol['Soluciones']:
        assert type(s) == np.ndarray or type(s) == str, f"La solución {s} no está en el formato correcto!" 

    # Numero de respuestas
    n_sol = len(sol['Soluciones'])
    n_bases = len(sol['Bases'])

    # Verificar número de soluciones y bases concuerda
    print(f"Verificando Solución {i}:")
    assert n_sol > 0, f"Llene el diccionario con soluciones!"
    assert n_sol == n_bases, "El número de bases y soluciones no concuerda! Ambos arreglos deben contener el mismo número de elementos"
    print("Formato OK! \n")
```

<!-- grade_id: cell-3bf33dfd1589b440, points: 15, grade: true, solution: false -->
```python
# Verificacion de soluciones (10pts)
maia.calificar_soluciones_basicas(Aa, ba, sol_c)
```

<!-- grade_id: cell-c5c43fe7d5009de5, points: 15, grade: true, solution: false -->
```python
# Verificacion de soluciones (20pts)
maia.calificar_soluciones_basicas(Ab, bb, sol_b)
```

<!-- grade_id: cell-b660ca0d3f10a0c7, points: 15, grade: true, solution: false -->
```python
# Verificacion de soluciones (20pts)
maia.calificar_soluciones_basicas(Ac, bc, sol_c)
```

# Problema 2: Automatización de las soluciones básicas

Hasta ahora el proceso para obtener estas soluciones ha sido manual y un poco repetitivo. No obstante, varios de estos pasos se pueden hacer de forma programática. La idea de este punto es poder automatizar el proceso de encontrar soluciones básicas de una matriz a partir de una función en Python:

Complete el bosquejo de la función dada a continuación (complete las tareas dadas) para encontrar todas las soluciones básicas del sistema de ecuaciones en el formato propuesto. Esta función recibe como parámetros de entrada la matriz $A$ y el vector $b$ de un sistema de ecuaciones lineal.

$$ A \mathbf{x} = \mathbf{b} \quad \longrightarrow \quad A \in \mathbb{R}^{m \times n}, \mathbf{x} \in \mathbb{R}^n, \mathbf{b} \in \mathbb{R}^m $$

Note que varias de estas lineas de código (así como algunas de las que usted debe completar) se han venido haciendo a lo largo del taller.

```python
# Función para hallar soluciones básicas

def soluciones_basicas(A, b, handle = True):
    """
    Encuentra las soluciones básicas de un sistema de ecuaciones de la forma:
      Ax = b

    Parámetros
    ----------
    A: np.array
      Matriz A (mxn)
    b: np.array
      Vector b (nx1)
    handle: bool
      Si manejar o no las submatrices singulares (reducción del subsistema eliminando la dependencia lineal)

    Retorna
    -------
    out: dict
      Diccionario con todas las soluciones básicas del sistema Ax=b en el formato establecido.

    """
    # Cambio de formato de vectores de entrada
    b = np.array(b).reshape(-1,1)
    A = np.array(A)

    # Guarda matriz original
    A_original = A.copy()

    # Transforma a forma canónica e identifica filas con pivote
    # Reduce la matriz para evitar dependencia lineal
    # (Será útil cuando se requiera reducir una submatriz)
    Ab = np.c_[A,b]
    _ , inds = sympy.Matrix(Ab).rref()
    inds = np.array(inds)    
    # Remueve filas sin pivote 
    A = A[inds]
    b = b[inds]
    
    # =====================================================
    # COMPLETAR ===========================================
    
    # Halla todas las posibles combinaciones de matrices cuadradas.
    # Hint: Revise las funciones importadas al inicio del notebook
    
    combinaciones = list()
    
    # Incializa la salida
    out = {'Soluciones':[], 'Bases': []}
    
    # =====================================================
    
    # Loop sobre todas las posibles combinaciones
    for comb in combinaciones:

        # Matriz de trabajo
        A_tmp = A[:,comb]
        
        # =====================================================
        # COMPLETAR ===========================================
        
        # Inicaliza vector solución en 0
        x_sol = None
        
        # Crea la base: Columnas fuera de la base en 0
        A_base = A_original.copy()
        A_base = None
        
        # =====================================================

        # Soluciona
        try:
            # =====================================================
            # COMPLETAR ===========================================
            
            # Soluciona para x hallando matriz inversa (solución única)
            
            x = None
            x.reshape(-1)
            
            # =====================================================

            # Guarda solución
            x_sol[np.array(comb)] = x
            out['Soluciones'].append(x_sol) 
            
            # Guarda base
            out['Bases'].append(A_base)
        
        # Caso que nos encontremos con una matriz singular (Rango deficiente)
        except Exception as ex:

            # Si se decide no manejar el error se guarda la excepcion (Matriz singular)
            if not handle:
                out['Soluciones'].append(str(ex))
                out['Bases'].append(A_base) 

            # Se intentan encontrar las soluciones de la submatriz
            else:
                try:
                    # Hallamos soluciones básicas de submatriz utilizando la misma función (Recurrencia)
                    # Recuerde que al inicio de la función se eliminan las filas linealmente dependientes
                    out2 = soluciones_basicas(A_tmp,b)

                    # Incorporamos en solución anterior
                    for i in range(out2.shape[0]):
                        # =====================================================
                        # COMPLETAR ===========================================
                        
                        # Se guarda la nueva solución
                        x_sol = None

                        # Se crea y se guarda la nueva A base
                        A_sol = A_base.copy()
                        A_sol = None

                        # Guarda las soluciones
                        out['Soluciones'].append(x_sol)
                        out['Bases'].append(A_sol)
                        
                        # =====================================================

                # En caso de que vuelva a fallar se guarda definitivamente el error
                except:
                    out['Soluciones'].append(str('Singular matrix'))
                    out['Bases'].append(A_base)

    # Formato dataframe
    out = pd.DataFrame(out)

    return out

```

Ahora pruebe la función que ha completado con las mismas matrices con las que se trabajo antes!

```python
# Matriz a)
sol = soluciones_basicas(Aa, ba)
maia.pretty_print(sol)
```

```python
# Matriz b)
sol = soluciones_basicas(Ab, bb, handle=True)
maia.pretty_print(sol)
```

```python
# Matriz c)
sol = soluciones_basicas(Ac, bc)
maia.pretty_print(sol)
```

## Calificación P2

<!-- grade_id: cell-cce23fdbd07b4c2c, points: 45, grade: true, solution: false -->
```python
maia.calificar_soluciones_basicas(Aa, ba, soluciones_basicas(Aa,ba))
maia.calificar_soluciones_basicas(Ab, bb, soluciones_basicas(Ab,bb))
maia.calificar_soluciones_basicas(Ac, bc, soluciones_basicas(Ac,bc))
```