![Banner](img/banner.png)

# **Actividad N°2:** Gram-Schimdt

***Matemáticas para Machine Learning***.

**Semana 1 - Lección 2:** Producto punto.

Fernando Enrique Lozano Martínez - Sergio David Salazar Isairias

## Introducción

### Descripción




El presente *jupyter notebook* contine todo el material para el desarrollo del taller 2 de la semana 1 del curso ***Matemáticas para Machine Learning***, correspondiente a la lección 2: Geometría de un espacio vectorial. En este se utilizará el conocimiento adquirido sobre producto interno para elaborar una función que realice el algoritmo de *Gram Schmidt*.

**Objetivos de Aprendizaje:**

*   Identificar las diferencias y similitudes entre conjuntos ortogonales y ortonormales.
*   Reconocer el funcionamiento del algoritmo *Gram Schmidt*.
*   Implementar el algoritmo de *Gram Schmidt*.



### Metodología

Para desarrollar el taller usted deberá editar las celdas de código dispuestas para esto. Estas estarán marcadas con el siguiente comentario:

```python
# =====================================================
# COMPLETAR ===========================================
# 

# =====================================================
```

Edite o complete el códgio dentro de estas lineas de comentarios. Dentro de estos comentarios encontrará indicaciónes de lo que debe hacer, así como algunas de las variables que debe utilizar o calcular (puede que estas tengan ya una estructura para llenar o esten solo igualadas a None, complete la asignación).

### Teoría



### Producto punto


Sea $V$ un espacio vectorial, el producto interno de $V$ es una función escalar.
\begin{equation*}
    \langle.,.\rangle : V^{2} \longrightarrow \mathbb{R}
\end{equation*}

El producto interno permite tener una intuición sobre "multiplicación de vectores". Existen diversos productos internos, pero en la clase solo se usara el producto punto. Dados los vectores $v= \{v_1, v_2, \dots v_n\}$ y $w=\{w_1, w_2, \dots, w_n \}$. El producto punto entre $v$ y $w$ denotado por $\langle v,w\rangle$, se define de la siguiente forma.
\begin{equation*}
    \langle v,w\rangle := v_{1}w_{1} + v_{2}w_{2} + \dots v_{n}w_{n}
\end{equation*}

El producto punto es de utilidad, pues permite tener una noción geométrica de $v$ y $w$. En particular, conocer si $v$ y $w$ son ortogonales. Lo anterior por medio de la siguiente formula.
\begin{equation*}
    \langle v,w\rangle = \|v \| \|w \| \cos{\phi}    
\end{equation*}

donde $\| .\|$ es la norma euclidiana y $\phi$ el ángulo entre $v$ y $w$. Si el ángulo entre $v$ y $w$ es de $90°$, $\cos{\phi}$ es cero, por tanto el producto punto es cero. De modo que el producto punto entre dos vectores es cero, si y solo si son ortogonales. 

Lo anterior facilita la comprensión del concepto de un *conjunto ortogonal*. Por definición un conjunto ortogonal es una colección de elementos, los cuales tienen la propiedad de que el angulo entre elementos diferentes es de 90°. Utilizando la operación producto punto se propone una nueva definición (que en escencia es la misma) la cual es que un conjunto es ortogonal si el producto punto entre elementos diferentes es cero.

```python
# importar librerias
import numpy as np
```

```python
# UTILIDADES =================================
# Correr una única vez por sesión ============
from maiautils import MaiaUtils
ipython = get_ipython()
mutils = MaiaUtils(ipython) 
# ============================================
```

```python
# Cambiar configuración de informe de errores
mutils.toggle_traceback()
```

#### Ejemplo - Versión a mano

Sea $n_1 = (1,2,0)$, $n_2 = (3,4,0)$, $n_3 = (-2,1,0)$.

*   $⟨n_1,n_2⟩=(1×3) + (2×4) + (0×0)= 11$.
*   $⟨n_1,n_3⟩=(1×-2) + (2×1) + (0×0)= 0$.
*   $⟨n_2,n_3⟩=(3×-2) + (4×1) + (0×0)= -2$.

Como se evidencia $n_1$ y $n_3$ con vectores ortogonales, pero el conjunto conformado por $\{n_{1},n_{2},n_{3}\}$ no es ortogonal.

#### Ejemplo - Versión en código


La libreria *numpy* provee la función *np.dot* la cual permite calcular el producto interno (también conocido como producto punto) entre dos *arrays*.

```python
# ===========================================================================
# VERSION SIMPLE
#   usando python como calculadora
# ===========================================================================

# DEFINIR n1, n2 y n3 
n1 = np.array([1,2,0])
n2 = np.array([3,4,0])
n3 = np.array([-2,1,0])

n1_dot_n2 = np.dot(n1,n2)
n1_dot_n3 = np.dot(n1,n3)
n2_dot_n3 = np.dot(n2,n3)

# IMPRIMIR EL RESULTADO
print(f'El producto punto entre n_1 y n_2 es igual a: {n1_dot_n2}')
print(f'El producto punto entre n_1 y n_3 es igual a: {n1_dot_n3}')
print(f'El producto punto entre n_2 y n_3 es igual a: {n2_dot_n3}')
```

```python
# ===========================================================================
# VERSION AVANZADA
#   usando recorridos de python
# ===========================================================================

# DEFINIR LISTA DE ARRAYs
n = [n1,n2,n3]

# CALCULAR LOS PRODUCTOS PUNTOS
# Definir el primer indice a revisar
i = 0
while i < len(n):
    
    # No se desea el producto punto entre un vector consigo mismo. Por tanto solo se revisan los vectores siguientes.
    j = i +1
    while j < len(n):
        
        # Calcula el producto interno entre el vector i-esimo y el vector j-esimo
        prod = np.dot(n[i],n[j])
        
        # Imprime el resultado de la operación
        print(f'El producto punto entre {n[i]} y {n[j]} es igual a : {prod}')
        
        # Avance del indice j
        j = j + 1
        
    # Avance del indice i    
    i = i + 1
```

![dot-product-example.png](img/dot-product-example.png)

En la imagen anterior se encuentra la gráfica de los vectores $n_{1}$, $n_{2}$ y $n_{3}$ en la curva de nivel $z=0$. Se observa que el angulo entre los $3$ elementos no es de $90°$, por lo que no corresponden a un conjunto ortogonal, lo cual concuerda con los resultados númericos.

### Gram Schmidt

El algoritmo de Gram-Schimdt es un procedimiento que permite generar un conjunto ortonormal a partir de un conjunto no ortonormal inicial. En particular, si se aplica el algoritmo a una base cualquiera, el resultado será una base ortonormal.


Dado un conjunto linealmente independiente $\beta = \{v_1, v_2, \dots, v_m\}$. Se calculan los elementos del nuevo conjunto $\alpha$ de la siguiente forma.

\begin{align*}
    \alpha_1 &= v_1\\
    \alpha_2 &= v_2 - \frac{\langle v_2,\alpha_{1}\rangle}{\langle\alpha_{1},\alpha_{1}\rangle} \alpha_{1}\\
    \alpha_3 &= v_3 - \frac{\langle v_3,\alpha_{1}\rangle}{\langle\alpha_{1},\alpha_{1}\rangle} \alpha_{1}- \frac{\langle v_3,\alpha_{2}\rangle}{\langle\alpha_{2},\alpha_{2}\rangle} \alpha_{2}\\
    &\vdots\\
    \alpha_{m} &= v_{m} - \sum_{i=1}^{m-1}\frac{\langle v_m,\alpha_{i}\rangle}{\langle\alpha_{i},\alpha_{i}\rangle} \alpha_{i}
\end{align*}


donde cada término $\frac{\langle v_i,\alpha_{j}\rangle}{\langle\alpha_{j},\alpha_{j}\rangle} \alpha_{j}$ se conoce como proyección de $v_i$ sobre $α_j$. 

Como resultado, se obtiene $\alpha = \{\alpha_1, \alpha_2, \dots, \alpha_m\}$ una base ortogonal. Ahora bien, para normalizar $\alpha$ y tener una base ortonormal, se divide cada $\alpha_i$ entre su norma.

\begin{equation*}
    \mathcal{E} = \left\{e_1, e_2,  \dots, e_m \right\}
\end{equation*}

donde $e_i = \frac{\alpha_i}{\|\alpha_i\|} $, con $\|α_i \|$ la norma euclidiana de $α_i$.

#### Versión a mano

Sea $\beta \subset \mathbb{R}^{3}$ linealmente independiente, calcule $\mathcal{E} \subset \mathbb{R}^{3}$ ortonormal por medio del algoritmo de Gram-Schmidt.

\begin{equation*}
    \beta = \left\{ 
    \begin{pmatrix}
    2\\
    0\\
    0
    \end{pmatrix},\begin{pmatrix}
    2\\
    4\\
    0
    \end{pmatrix},\begin{pmatrix}
    0\\
    2\\
    2
    \end{pmatrix}
    \right\}
\end{equation*}

Declarar $\alpha_1$.
\begin{equation*}
    \alpha_1 = \begin{pmatrix}
    2\\
    0\\
    0
    \end{pmatrix}
\end{equation*}

Calcular $\alpha_2$.
\begin{equation*}
    \alpha_2 = \begin{pmatrix}
    2\\
    4\\
    0
    \end{pmatrix} - \frac{4}{4}\begin{pmatrix}
    2\\
    0\\
    0
    \end{pmatrix}
    = \begin{pmatrix}
    0\\
    4\\
    0
    \end{pmatrix}
\end{equation*}

Calcular $\alpha_3$.
\begin{equation*}
    \alpha_3 = \begin{pmatrix}
    0\\
    2\\
    2
    \end{pmatrix} - \frac{0}{4}\begin{pmatrix}
    2\\
    0\\
    0
    \end{pmatrix} - \frac{8}{16}\begin{pmatrix}
    0\\
    4\\
    0
    \end{pmatrix}
    = 
    \begin{pmatrix}
    0\\
    0\\
    2
    \end{pmatrix}
\end{equation*}

Dado $\alpha$, se calcula $\mathcal{E}$.
\begin{equation*}
    \mathcal{E} = \left\{\begin{pmatrix}
    1\\
    0\\
    0
    \end{pmatrix},\begin{pmatrix}
    0\\
    1\\
    0
    \end{pmatrix},\begin{pmatrix}
    0\\
    0\\
    1
    \end{pmatrix} \right\}
\end{equation*}

#### Versión en código

```python
# DEFINIR LOS VECTORES DEL CONJUNTO BETA
b1 = np.array([2,0,0])
b2 = np.array([2,4,0])
b3 = np.array([0,2,2])

a1 = b1

# =====================================================
# COMPLETAR ===========================================
# -
# reemplace la lista de ceros por la expresión correcta para a2
# -

a2 = [0,0,0]

# =====================================================

a3 = b3 - (np.dot(b3,a1)/np.dot(a1,a1))*a1 - (np.dot(b3,a2)/np.dot(a2,a2))*a2

e1 = (1/np.linalg.norm(a1))*a1
e2 = (1/np.linalg.norm(a2))*a2

# =====================================================
# COMPLETAR ===========================================
# -
# reemplace la lista de ceros por la expresión correcta para e3
# -

e3 = [0,0,0]

# =====================================================

print(f'Los elementos del conjunto ortonormal resultante de realizar el algoritmo de Gram - Schmidt son: {e1}, {e2}, {e3}')
```

![gram-s-example.png](img/gram-s-example.png)

En la imagen anterior se encuentra el algoritmo de Gram-Schmidt aplicado a un conjunto en $\mathbb{R}^2$. Donde $v_1 = (1,1)$ y $v_2 = (-1,2)$.

### Punto 1. Conjuntos ortogonales

Escriba una función en Python que tenga como párametro de entrada una lista con los vectores un conjunto dado. La función debe retornar un booleano que índique si el conjunto es ortogonal o no.

Luego, utilice la función para determinar si los siguientes conjuntos son ortogonales.


*   $β = \{(0,1,1),(3,4,9),(10,0,1)\}$
*   $β = \{(0,1,1),(3,0,0),(0,-1,1)\}$






```python
def es_ortogonal(v_list:list):
    """ 
    Determina si un conjunto dado es ortogonal
    ___________________________________
    Entrada:
    v_list:    [list] Conjunto de vectores
    ___________________________________
    Salida:
    ortogonal: [boolean] True si el conjunto es ortogonal, False si no es ortogonal
    """

    # Variable booleana que se retorna al final del código
    ortogonal = True

    # =====================================================
    # COMPLETAR ===========================================
    # -
    # HINT: piense en un algoritmo que le permita calcular
    #      todas las combinaciones de productos punto
    # =====================================================
    
    return ortogonal
```

<!-- grade_id: cell-4463c9519c3e0d52, points: 0, grade: true, solution: false -->
```python
# PROBAR =======================================================================================
v = [[0,1,1],[3,4,9],[10,0,1]]
ortho = es_ortogonal(v)
print(ortho)
### BEGIN HIDDEN TESTS
### END HIDDEN TESTS
```

<!-- grade_id: cell-7ceb318ccf09d5b5, points: 20, grade: true, solution: false -->
```python
# PROBAR =======================================================================================
v = [[0,1,1],[3,0,0],[0,-1,1]]
ortho = es_ortogonal(v)
print(ortho)
### BEGIN HIDDEN TESTS
### Pruebas básicas

# DICCIONARIO CON LOS ELEMENTOS QUE SE USARAN PARA COMPROBAR LA FUNCION

dic_assert = {'PRUEBA':[],'BOOLEAN':[],'MENSAJE':[]}

dic_assert['PRUEBA'].append([[1,2],[-2,1]])
dic_assert['BOOLEAN'].append(True)
dic_assert['MENSAJE'].append("Vectores: [1,2], [-2,1] \nValor esperado: son ortogonales.")

dic_assert['PRUEBA'].append([[0,0],[-2,1]])
dic_assert['BOOLEAN'].append(True)
dic_assert['MENSAJE'].append("Vectores: [0,0], [-2,1] \nValor esperado: son ortogonales \nNota: El vector cero es ortogonal a cualquier otro vector.")

dic_assert['PRUEBA'].append([[1,1],[-2,1]])
dic_assert['BOOLEAN'].append(False)
dic_assert['MENSAJE'].append("Vectores: [1,1], [-2,1] \nValor esperado: NO son ortogonales")

dic_assert['PRUEBA'].append([[1,1],[-2,1],[3,4]])
dic_assert['BOOLEAN'].append(False)
dic_assert['MENSAJE'].append("Vectores: [1,1], [-2,1], [3,4] \nValor esperado: NO son ortogonales \nNota: Dado que hay 3 vectores no nulos en un espacio de dimensión 2, el conjunto conformado por estos no puede ser ortogonal.")

dic_assert['PRUEBA'].append([[2,0,0],[0,0,9],[0,3,0]])
dic_assert['BOOLEAN'].append(True)
dic_assert['MENSAJE'].append("Vectores: [2,0,0], [0,0,9], [0,3,0] \nValor esperado: son ortogonales")

dic_assert['PRUEBA'].append([[0,0],[0,0],[0,1]])
dic_assert['BOOLEAN'].append(True)
dic_assert['MENSAJE'].append("Vectores: [0,0], [0,0], [0,1] \nValor esperado: son ortogonales")


for i in range(0,len(dic_assert['PRUEBA'])):
    """Recorre el diccionario dic_assert realizando una comparación entre la salida del código y su valor esperado""" 
    prueba = dic_assert['PRUEBA'][i]
    value = dic_assert['BOOLEAN'][i]
    resumen = dic_assert['MENSAJE'][i]
    assert es_ortogonal(prueba) == value, f"\nRESPUESTA INCORRECTA EN LA PRUEBA {i+1} \nRESUMEN: \n{resumen}."
### END HIDDEN TESTS
```

### Punto 2. Gram - Schmidt

Escriba una función en Python que realiza el algoritmo de Gram-Schmidt. 

*Nota*: Si el conjunto que se utiliza para el algoritmo de Gram-Schmidt es linealmente dependiente, se produce el vector cero en una de las iteraciones, lo cual genera inconvientes en el calculó de la siguiente iteración, pues su norma es cero.

```python
def gram_schmindt(v_list:list):
    """
    Recibe un conjunto de vectores linealmente independientes 
      y obtiene un conjunto ortonormal por medio del algoritmo de gram-schimdt
    ___________________________________
    Entrada:
    v_list: [list] Lista de los elementos del conjunto \beta'
    ___________________________________
    Salida:
    v_list_gs: [list]  Lista de los vectores resultantes del algoritmo de gram-schimdt
    """

    # Lista que contiene los elementos del conjunto ortonormal final.
    v_list_gs = []

    # Se define el vector de ceros.
    vector_cero = np.zeros(len(v_list[0]))

    # Recorrido de los elementos de v_list
    for i in range(0,len(v_list)): 
        # Elemento actual de la iteración (v_i).
        actual = np.array(v_list[i])

        # Vector que contiene la suma de las proyecciones de 'actual' sobre los elementos de 'v_list_gs'
        proyeccion = vector_cero

        # Rcorrido de los elementos del conjunto ortogonal
        for j in range(0,len(v_list_gs)):
            # Elemento del conjunto ortogonal parcial
            aj = np.array(v_list_gs[j])

            # =====================================================
            # COMPLETAR ===========================================
            
            # Actualización del vector 'proyeccion'
            # cambie vector_cero por la expresión corrrecta
            # recuerde que las proyecciones se 'acumulan'
            
            proyeccion = vector_cero
            
            # =====================================================
            
            
        # =====================================================
        # COMPLETAR ===========================================
        
        # Calcular el siguiente vector del conjunto ortonormal
        # cambie vector_cero por la expresión corrrecta
        
        a_siguiente = vector_cero
        
        # =====================================================

        # Determina si el vector resultante de la iteración (i) es el vector cero.
        if (a_siguiente!=vector_cero).any():
            # =====================================================
            # COMPLETAR ===========================================

            # normalice el vector y agregue el resultado a v_list_gs
            2+2
            # =====================================================

    # Transformar cada elemento de v_list_gs de formato array a formato de lista (list).
    for i in range(len(v_list_gs)):
        v_list_gs[i] = v_list_gs[i].tolist()

    # Retornar la lista con los elementos resultantes del algoritmo de Gram-Schmidt
    return v_list_gs
```

```python
# PROBAR =======================================================================================
v = [[2,0,0],[2,4,0],[0,2,2]]
v_gs = gram_schmindt(v)
print(v_gs)
```

<!-- grade_id: cell-c1e16eeb804d6085, points: 80, grade: true, solution: false -->
```python
# PROBAR =======================================================================================
v = [[1,1],[-1,2]]
v_gs = gram_schmindt(v)
print(v_gs)
### BEGIN HIDDEN TESTS

# Solución ========================================

# Prueba ========================================================================
dic_assert = {'PRUEBA':[],'VALOR_ESPERADO':[]}

dic_assert['PRUEBA'].append([[0,0,0]])
dic_assert['VALOR_ESPERADO'].append([])

dic_assert['PRUEBA'].append([[2,0],[-1,2]])
dic_assert['VALOR_ESPERADO'].append([[1.0, 0.0], [0.0, 1.0]])

dic_assert['PRUEBA'].append([[0,0],[3,4]])
dic_assert['VALOR_ESPERADO'].append([[0.6, 0.8]])

dic_assert['PRUEBA'].append([[2,0,0],[2,4,0],[0,2,2],[4,6,2]])
dic_assert['VALOR_ESPERADO'].append([[1.0, 0.0, 0.0], [0.0, 1.0, 0.0], [0.0, 0.0, 1.0]])

dic_assert['PRUEBA'].append([[1,1,1,1],[2,2,2,2],[3,3,3,3]])
dic_assert['VALOR_ESPERADO'].append([[0.5, 0.5, 0.5, 0.5]])

dic_assert['PRUEBA'].append([[1,0,0,0],[0,4,3,0],[0,0,0,2]])
dic_assert['VALOR_ESPERADO'].append([[1.0, 0.0, 0.0, 0.0], [0.0, 0.8, 0.6, 0.0], [0.0, 0.0, 0.0, 1.0]])

dic_assert['PRUEBA'].append([[1,0,0,0],[0,4,3,0],[0,0,0,2],[0,4,3,2],[0,0,0,0],[1,4,3,2]])
dic_assert['VALOR_ESPERADO'].append([[1.0, 0.0, 0.0, 0.0], [0.0, 0.8, 0.6, 0.0], [0.0, 0.0, 0.0, 1.0]])


for i in range(0,len(dic_assert['PRUEBA'])):
    """ Recorre el diccionario dic_assert realizando una comparación entre la salida del código y su valor esperado""" 
    prueba = dic_assert['PRUEBA'][i]
    value = dic_assert['VALOR_ESPERADO'][i]
    assert (gram_schmindt(prueba) == value), f"\nRESPUESTA INCORRECTA EN LA PRUEBA {i+1} \nEL VALOR ESPERADO ES \n{value}"
### END HIDDEN TESTS
```